# Assignment 3: Informed Search

##  Grid Navigation
In this section, you will investigate the problem of navigation on a two dimensional grid with obstacles. The goal is to produce the shortest path between a provided pair of points, taking care to maneuver around the obstacles as needed. Path length is measured in Euclidean distance. Valid directions of movement include up, down, left, right, up-left, up-right, down-left, and down-right. 

### Note:
You are expected to write code where you see **your code here**.  
Make sure you delete the lines with **pass** and **raise NotImplementedError** or your code may not run correctly.

### Task
Your task is to write a function find_path(start, goal, scene) which returns the shortest path from the start point to the goal point that avoids traveling through the obstacles in the grid. 

### Structure/Representation
For this problem, points will be represented as two-element tuples of the form (row, column), and scenes will be represented as two-dimensional lists of Boolean values, with False values corresponding empty spaces and True values corresponding to obstacles. 

### Output
Your output should be the list of points in the path, and should explicitly include both the start point and the goal point. If multiple optimal solutions exist, any of them may be returned. If no optimal solutions exist, or if the start point or goal point lies on an obstacle, you should return the sentinal value None. If start and goal state are the same, return None.

### Implementation
Your implementation should consist of 
* an A* search 
* the straight-line Euclidean distance heuristic. 

Helper functions are allowed and encouraged.

## 1. Euclidean distance
First, let's define a function used for euclidean distance.

<!-- grade_id: cell-solution-1, points: 0, grade: false, solution: true -->
```python
############################################################
# Section 1 Grid Navigation - Euclidiean distance
############################################################
def grid_euclidean_distance(current_state, goal_state):
    #pass
    
    # your code here
    #raise NotImplementedError
    from math import sqrt
    return sqrt(sum([(x - y)**2 for (x, y) in zip(current_state, goal_state)]))
```

<!-- grade_id: cell-test-1, points: 25, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

current_state1, goal_state1 = (1,1), (1,1)
assert grid_euclidean_distance(current_state1, goal_state1) == 0

current_state2, goal_state2 = (1,2), (1,1)
assert grid_euclidean_distance(current_state2, goal_state2) == 1

current_state3, goal_state3 = (3,2), (1,2)
assert grid_euclidean_distance(current_state3, goal_state3) == 2

current_state4, goal_state4 = (1,1), (2,2)
assert grid_euclidean_distance(current_state4, goal_state4) == 2**0.5
print("test passed!")
```

Output:
```
test passed!
```

## 2. Helper functions
Next, let's define a functions that finds the successors.

<!-- grade_id: cell-solution-2, points: 0, grade: false, solution: true -->
```python
# your code here
#raise NotImplementedError

def grid_successors(current, scene):
    # assuming current is not at a false position
    # your code here
    #raise NotImplementedError
    x, y = current
    return tuple([(x + dx, y + dy) for dx in range(-1,2) for dy in range(-1,2) \
                  if x + dx >= 0 and y + dy >= 0 and x + dx < len(scene) and y + dy < len(scene[0]) \
            and not scene[x+dx][y+dy] and not (dx == 0 and dy == 0)])[::-1]
```

<!-- grade_id: cell-test-2, points: 25, grade: true, solution: false -->
```python
scene1 = [[True, True, True],
         [False, False, True]]
assert grid_successors((1, 2), scene1) == ((1, 1),)
assert grid_successors((0, 1), scene1) == ((1, 1),(1, 0),)
print("test passed!")
```

Output:
```
test passed!
```

## 3. Find path
Finally let's implement the path search.

<!-- grade_id: cell-solution-3, points: 0, grade: false, solution: true -->
```python
############################################################
# Grid Navigation
############################################################
import collections, itertools, queue, random, copy

 
def find_path(start, goal, scene):
    #pass

    # your code here
    #raise NotImplementedError
    visited = collections.defaultdict(bool)
    d = collections.defaultdict(int)
    p = {}
    q = queue.PriorityQueue()
    q.put((0, start))
    d[start] = 0
    while q.qsize() > 0:
        _, cur = q.get()
        visited[cur] = True
        if cur == goal:
            break
        for child in grid_successors(cur, scene):
            if not visited[child]:
                p[child] = cur
                d[child] = d[cur] + 1
                q.put((d[child] + grid_euclidean_distance(child, goal), child))
    if not goal in p:
        return None
    path = []
    cur = goal
    while cur != start:
        path = [cur] + path
        cur = p[cur]
    path = [start] + path
    return path
```

<!-- grade_id: cell-test-3, points: 50, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################
scene1 = [[False, False, False], 
          [False, True , False], 
          [False, False, False]] 

assert find_path((0, 0), (2, 1), scene1) == [(0, 0), (1, 0), (2, 1)] 

scene2 = [[False, True, False], 
          [False, True, False], 
          [False, True, False]] 
assert find_path((0, 0), (0, 2), scene2) is None
print("test passed!")
```

Output:
```
test passed!
```

```python

```